<problem display_name="Randomized Vector Addition">
<script type="loncapa/python">
<![CDATA[
seed = random.random()
html_file= "/static/RandomizedVectorProjection.html?seed=SEED".replace("SEED",str(seed))
import vectordraw
grader = vectordraw.Grader()
]]>
</script>
<p>Draw:</p>
<ul>
    <li>\(B_1\), the projection of B onto A</li>
    <li>\(B_2\), the projection of B perpendicular to A</li>
</ul>
<customresponse cfn="grader.cfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state="{a:1}"
    width="800"
    height="480"
    html_file ="$html_file"
    />
</customresponse>

<p>Random seed is $seed.</p>

<p>Ideas:</p>
<ul>
    <li>This should be preceeded by a version in which the angle between A and B is always between 0 and 90.</li>
    <li>I'd like to do this as a tangential/radial acceleration problem</li>
</ul>

</problem>

