<problem display_name="Randomized Vector Addition">
<script type="loncapa/python">
<![CDATA[
import vectordraw
grader = vectordraw.Grader()
]]>
</script>
<p>Draw A+B.</p>
<customresponse cfn="grader.cfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="800"
    height="480"
    html_file="/static/RandomizedVectorDraw.html"
    />
</customresponse>
</problem>

