<problem display_name="Draw Force Diagram">
<script type="loncapa/python">
<![CDATA[

import vectordraw

### Custom Checks ###	
custom_checks = {}

success_message = 'Good job!'

grader = vectordraw.Grader(success_message, custom_checks)
]]>
</script>

<p>Below is a graph of the curve defined implicitly by \(y^4+xy=x^3-x+2\).</p>

<p>Determine the slope of this curve at the point in the 1st quadrant whose \(x\)-value is \(x=1\), then the line tangent to the curve at this point.</p>

<p><em>Orient your line by dragging its handles; shift your tangent line by dragging the line's center. You may need to try this a few times to get the appropriate angle.</em></p>


<customresponse cfn="grader.cfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="800"
    height="480"
    html_file="/static/TangentLine_A.html"
    />
</customresponse>

<p><strong>Answer:</strong> A line of slope \(\frac{1}{5}\) touching the point \((1,1)\)</p>
<p><strong>Grader:</strong> Uses two points.</p>

</problem>





