<problem display_name="Randomized Vector Addition">
<script type="loncapa/python">
<![CDATA[
seed = random.random()
html_file= "/static/RandomizedVectorAddition.html?seed=SEED".replace("SEED",str(seed))
import vectordraw
grader = vectordraw.Grader()
]]>
</script>
<p>Draw A+B.</p>
<customresponse cfn="grader.cfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state="{a:1}"
    width="800"
    height="480"
    html_file ="$html_file"
    />
</customresponse>

<p>Random seed is $seed.</p>

</problem>

